# BLU01 - Exercises Notebook

**Important Note**

Attention Windows users: The grader will run on Linux, and using power shell statements will not work.

```python
import chardet
import hashlib # for grading purposes
import math
import numpy as np
import csv
import pandas as pd
import random
```

## Q1: Use a shell command to count lines





<!-- grade_id: cell-05a3d2245d57305d, points: 0, grade: false, solution: true -->
```python
# Count the total number of lines of files data/exercises/surfing.txt and data/exercises/body_boarding.txt
# and store the result in variable count_total.  
# There is more than one way to do this. Hint: concatenate both files.
count_total_1 = ! wc -l < data/exercises/surfing.txt
count_total_2= ! wc -l < data/exercises/body_boarding.txt
count_total= int(count_total_1[0])+int(count_total_2[0])
count_total
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
44
```

<!-- grade_id: cell-2e69c6137e73d90c, points: 2, grade: true, solution: false -->
```python
expected_hash = '71ee45a3c0db9a9865f7313dd3372cf60dca6479d46261f3542eb9346e4a04d6'
assert hashlib.sha256(str(count_total).encode()).hexdigest() == expected_hash, "count_total is wrong"
```

Now store in 2 variables the first third and the last third of the lines existent in file data/exercises/surfing.txt.

<!-- grade_id: cell-14cf2619322a8763, points: 0, grade: false, solution: true -->
```python
# NOT GRADED optional exercise!!
# Add the first third and the last third of the lines to the variables first_third and last_third. 
# Make it work to any files using count_total_surfing/3 for instance in the bash commands
count_total_surfing = ! wc -l < data/exercises/surfing.txt
f_3rd=int(count_total_surfing[0])/3
first_third = ! head -4 data/exercises/surfing.txt
last_third = ! tail -4 data/exercises/surfing.txt
# YOUR CODE HERE
#raise NotImplementedError()
```

```python
first_third
```

Output:
```
['Surfing is a surface water pastime in which the wave rider, referred to as a surfer, ',
 'rides on the forward part, or face, of a moving wave, which usually carries the ',
 'surfer towards the shore. Waves suitable for surfing are primarily found in the ',
 'ocean, but can also be found in lakes or rivers in the form of a standing wave or ']
```

<!-- grade_id: cell-20d96111430874d4, points: 0, grade: true, solution: false -->
```python
#Test the first_third and last_third exercise by running the following (ungraded) asserts.
assert first_third[-1][0] == 'o'
assert last_third[0][0] == 'o'
```

## Q2: Read a file with specific delimiter

Read file **data/exercises/surfing_sessions.csv** into a pandas DataFrame.

First, you should preview the file using a shell command in order to find out the used delimiter, and other properties of this file.

Then, you should use function read_csv to read the data into a DataFrame. The resulting DataFrame should have the last column as index.

<!-- grade_id: cell-942e0590467e20d9, points: 0, grade: false, solution: true -->
```python
# Use a shell command to preview the data
! head -2 data/exercises/surfing_sessions.csv
# YOUR CODE HERE
#raise NotImplementedError()

# Use function read_csv to read the data into a DataFrame
df2 = pd.read_csv('data/exercises/surfing_sessions.csv', sep='|', index_col=-1)
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
surfer|wave_power|location|visit_id
3|37,6|Rocky Point|146097
```

<!-- grade_id: cell-0c275c9acaa9c74d, points: 2, grade: true, solution: false -->
```python
assert df2.loc[126903, 'location'] == 'Pupukea', "df2 data is wrong"
assert set(df2.columns) == {'location', 'surfer', 'wave_power'}, "df2 data is wrong"
assert len(df2) == 108, "df2 data is wrong"
assert df2.index[0] == 146097, "df2 data is wrong"

expected_hash = 'a27dd8b953f484c0d74059da102e4b8e513630292de5d17ff5585fc6743d62f0'
assert hashlib.sha256(df2.index.name.encode()).hexdigest() == expected_hash
```

## Q3: Read a csv file with problems

Read file **data/exercises/surfing_sessions_w_problems.csv** using function `read_csv`. Pay attention to the following:
* remove all lines where the number of columns is different from the majority of lines.
* use the first column as index
* there are some inputs in the file that should be interpreted as NaN, make sure you select the right one when reading the file!

<!-- grade_id: cell-2771b707556c08d2, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/surfing_sessions_w_problems.csv with read_csv
! tail -30 data/exercises/surfing_sessions_w_problems.csv
fields=['visit_id','surfer','wave_power','location']
df3 = pd.read_csv('data/exercises/surfing_sessions_w_problems.csv', sep=';', index_col=0, usecols=fields,
                 na_values= {'wave_power':['missing-data', '-1']})
df3=df3.drop([126902, 119735,126340,119736], axis=0)
#! wc -l < data/exercises/surfing_sessions_w_problems.csv
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
visit_id;surfer;wave_power;location
146097;3;37.6;Rocky Point
146087;3;5.7;Chuns
130395;3;-1;Noronha(Cacimba)
130363;3;5;Rocky Point
130362;3;11.3;Noronha(Conceicao)
130360;3;8.9;Noronha(Boldo)
126903;3;32.7;Pupuke
126902;3;missing-data;Pupukea;;
126591;3;9.7;Ehukai
126590;3;10.4;Changes
126341;3;8.2;Changes
126340;3;38.6;Laniakea;
125540;3;34.1;Laniakea
122199;3;14.4;Glass Doors
121170;3;34.9;Walter’s West
120454;3;-1;Gas Chambers
119736;3;21.7;Rocky Point;
119735;3;21.2;Rocky Point;;
```

```python
#df3.index
df3['wave_power'].mean()
```

Output:
```
17.741666666666667
```

```python
df3.dtypes
```

Output:
```
surfer          int64
wave_power    float64
location       object
dtype: object
```

<!-- grade_id: cell-601b354802964776, points: 2, grade: true, solution: false -->
```python
assert np.isnan(df3.loc[130395, 'wave_power']), "df3 data is wrong"
assert df3.loc[126903, 'wave_power'] == 32.7, "df3  data is wrong"

mean_wave_power = df3['wave_power'].mean()
assert math.isclose(17.741, mean_wave_power, rel_tol=1e-3), "df3 data is wrong"

expected_hash = 'a27dd8b953f484c0d74059da102e4b8e513630292de5d17ff5585fc6743d62f0'
assert hashlib.sha256(df3.index.name.encode()).hexdigest() == expected_hash
```

## Q4: Repair a csv file after importing
Read the same file **data/exercises/surfing_sessions_w_problems.csv** using function `csv.reader()`. But now, be sure you don't miss any lines with relevant information! 

* use csv module to import everything to a list of lists
* create a df with only 4 meaningful columns and then 
* replace garbage values with NaN's 
* format the columns with the right type
* set the index of the df to column `visit_id`

<!-- grade_id: cell-119041388e17fb72, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/surfing_sessions_w_problems.csv using csv.reader() and add result to
f = open('data/exercises/surfing_sessions_w_problems.csv', 'r')

# read the csv into a list of lists
lines = list(csv.reader(f, delimiter= ';'))
lines = [i[:4] for i in lines]
f.close() 
# YOUR CODE HERE
#raise NotImplementedError()
# create a dataframe using the line list with only 4 columns
df4 = pd.DataFrame(lines[1:], columns=lines[0])
# YOUR CODE HERE
#raise NotImplementedError()
# replace invalid values with nan (np.nan)
df4['wave_power'] = df4['wave_power'].replace({'missing-data':np.nan, '-1': np.nan})
# YOUR CODE HERE
#raise NotImplementedError()
# set types per dataframe column (always use int64 when int's are needed)
df4['visit_id'] = df4['visit_id'].astype('int64')
df4['surfer'] = df4['surfer'].astype('int64')
df4['wave_power'] = df4['wave_power'].astype('float64')
df4['location'] = df4['location'].astype('object')
#raise NotImplementedError()
# set a new index to the dataframe
df4=df4.set_index('visit_id')
#raise NotImplementedError()


```

```python
df4.head(20)
```

Output:
```
surfer  wave_power            location
visit_id                                        
146097         3        37.6         Rocky Point
146087         3         5.7               Chuns
130395         3         NaN    Noronha(Cacimba)
130363         3         5.0         Rocky Point
130362         3        11.3  Noronha(Conceicao)
130360         3         8.9      Noronha(Boldo)
126903         3        32.7              Pupuke
126902         3         NaN             Pupukea
126591         3         9.7              Ehukai
126590         3        10.4             Changes
126341         3         8.2             Changes
126340         3        38.6            Laniakea
125540         3        34.1            Laniakea
122199         3        14.4         Glass Doors
121170         3        34.9       Walter’s West
120454         3         NaN        Gas Chambers
119736         3        21.7         Rocky Point
119735         3        21.2         Rocky Point
```

<!-- grade_id: cell-a42299f9047590d5, points: 2, grade: true, solution: false -->
```python
assert np.isnan(df4.loc[130395, 'wave_power']), "df4 content is wrong"
assert df4.loc[126340, 'wave_power'] == 38.6, "df4 content is wrong"

mean_wave_power = df4['wave_power'].mean()
assert math.isclose(19.626, mean_wave_power, rel_tol=1e-3), "df4 content is wrong"

assert df4['surfer'].dtype == np.int64, "df4 types are wrong"
assert df4['location'].dtype == object, "df4 types are wrong"
assert df4['wave_power'].dtype == float, "df4 types are wrong"
assert df4.index.dtype == np.int64, "df4 types are wrong"


expected_hash = 'a27dd8b953f484c0d74059da102e4b8e513630292de5d17ff5585fc6743d62f0'
assert hashlib.sha256(df4.index.name.encode()).hexdigest() == expected_hash
```

## Q5: Read a JSON file

Read file **data/exercises/portugal_production_of_electricity_gwh.json**.

<!-- grade_id: cell-5c2125479f7e50e5, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_production_of_electricity_gwh.json with read_json
#! cat data/exercises/portugal_production_of_electricity_gwh.json
df5 = pd.read_json('data/exercises/portugal_production_of_electricity_gwh.json', orient='table')
#df5.head()
# YOUR CODE HERE
#raise NotImplementedError()
```

```python
df5
```

Output:
```
Year  Total  Total renewables  Hydropower > 10MW  Hydropower < 10MW  \
0   1995  33264              9501               7962                492   
1   1996  34520             15895              14007                658   
2   1997  34207             14301              12537                638   
3   1998  38984             14224              12488                566   
4   1999  43287              8915               7042                589   
5   2000  43764             13260              11040                675   
6   2001  46509             16083              13605                770   
7   2002  46107             10190               7551                706   
8   2003  46852             18037              15163                891   
9   2004  45105             12597               9570                577   
10  2005  46575              8616               4737                381   
11  2006  49041             16187              10633                834   
12  2007  47253             1
```

<!-- grade_id: cell-3f085fa2cdad9319, points: 2, grade: true, solution: false -->
```python
assert len(df5) == 23
assert set(df5.columns) == {'Biomass', 'Geothermal power', 'Hydropower < 10MW', 'Hydropower > 10MW', 
                           'Photovoltaic', 'Total','Total renewables','Windpower','Year'}, "df5 columns are wrong"

expected_hash = '48beaa4bb16f0656bb4e4d3abb5da6ff6a50b52be037354f8c33455a7534102e'
assert hashlib.sha256(str(df5.loc[:,'Hydropower > 10MW'].sum()).encode()).hexdigest() == expected_hash, "df5 content is wrong"
```

## Q6: Read an Excel file

Read file **data/exercises/portugal_gas_emissions_per_year.xlsx** using function read_excel. Pay attention to the following:

* you should grab the table "Series" in sheet "metadata"
* use column 'Serie' as index
* make sure you keep only the rows and columns with data
* set the variable distinct_scales with the number of ... distinct scales found in the table

```python
#! pip install xlrd==1.2.0
#! pip install lxml==4.5.2
#! pip install openpyxl==3
```

Output:
```
Collecting xlrd==1.2.0
  Downloading xlrd-1.2.0-py2.py3-none-any.whl (103 kB)
     |████████████████████████████████| 103 kB 3.1 MB/s            
[?25hInstalling collected packages: xlrd
  Attempting uninstall: xlrd
    Found existing installation: xlrd 2.0.1
    Uninstalling xlrd-2.0.1:
      Successfully uninstalled xlrd-2.0.1
Successfully installed xlrd-1.2.0
```

<!-- grade_id: cell-017a7c31b0ddc38e, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_gas_emissions_per_year.xlsx with read_excel
import os
df6 = pd.read_excel(os.path.join('data', 'exercises','portugal_gas_emissions_per_year.xlsx'), sheet_name="metadata")
df6
# YOUR CODE HERE
#raise NotImplementedError()

#distinct_scales = ...
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
~/.virtualenvs/slu00/lib/python3.6/site-packages/pandas/io/excel/_xlrd.py in load_workbook(self, filepath_or_buffer)
     34             return open_workbook(file_contents=data)
     35         else:
---> 36             return open_workbook(filepath_or_buffer)
     37 
     38     @property

~/.virtualenvs/slu00/lib/python3.6/site-packages/xlrd/__init__.py in open_workbook(filename, logfile, verbosity, use_mmap, file_contents, encoding_override, formatting_info, on_demand, ragged_rows, ignore_workbook_corruption)
    168     :param unnumbered: If true, omit offsets (for meaningful diffs).
    169     """
--> 170     from .biffh import biff_dump
    171     bk = Book()
    172     bk.biff2_8_load(filename=filename, logfile=outfile, )

XLRDError: Excel xlsx file; not supported
```

<!-- grade_id: cell-91d5f07e40585680, points: 2, grade: true, solution: false -->
```python
assert distinct_scales == 2
assert isinstance(df6, pd.DataFrame)
expected_hash = '60c3ad36f77e7366103fe36a3f551a0cde7d64e26d3102ddb8b953d6208a6006'
assert hashlib.sha256(
        df6.loc[
            df6.index=="Nitrogen oxides", 
            "Measure Unit"][0].encode()
    ).hexdigest() == expected_hash, "df6 is wrong"


```

## Q7: Find the encoding of a file

Find the encoding used in file **data/exercises/cities.csv**, using the method that was shown in the Learning Units.

Then, read the data into a DataFrame, using the read_csv method and find the `City` characters that has distance equal to 2377.

<!-- grade_id: cell-bd6944fd63bb38d8, points: 0, grade: false, solution: true -->
```python
# Find the encoding of file data/exercises/cities.csv
#encoding = 'IBM866'
#chardet.detect(open(os.path.join('data', 'exercises', 'cities.csv'), 'rb').read())
# YOUR CODE HERE
#raise NotImplementedError()
# Read the file into a DataFrame
df7 = pd.read_csv(os.path.join('data', 'exercises', 'cities.csv'), encoding='IBM866')
df7
# YOUR CODE HERE
#raise NotImplementedError()

# Find the name of the city with distance = 2377
city_found = df7.City[df7.distance==2377] #(assign a string)
city_found
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
6    Агидель
Name: City, dtype: object
```

<!-- grade_id: cell-687fb83c97e03355, points: 2, grade: true, solution: false -->
```python
assert isinstance(df7, pd.DataFrame)

expected_hash_1 = '969aef39a1d4cb1c5928c774cd7a4e3ccfc064a18fbd43a70193a2631d8a122d'
assert hashlib.sha256(encoding.encode()).hexdigest() == expected_hash_1, "encoding is wrong"

expected_hash_2 = '22522b6940ec2bcc42ea98f55878bee105c27254c4340fbce9fbc951b5aed078'
assert hashlib.sha256(city_found.encode()).hexdigest() == expected_hash_2, "city_found is wrong"
```

Output:
```
<ipython-input-123-dc5bac29be7e> in <module>
      5 
      6 expected_hash_2 = '22522b6940ec2bcc42ea98f55878bee105c27254c4340fbce9fbc951b5aed078'
----> 7 assert hashlib.sha256(city_found.encode()).hexdigest() == expected_hash_2, "city_found is wrong"

~/.virtualenvs/slu00/lib/python3.6/site-packages/pandas/core/generic.py in __getattr__(self, name)
   5272             if self._info_axis._can_hold_identifiers_and_holds_name(name):
   5273                 return self[name]
-> 5274             return object.__getattribute__(self, name)
   5275 
   5276     def __setattr__(self, name: str, value) -> None:

AttributeError: 'Series' object has no attribute 'encode'
```

## Q8: Import a Random Sample of a Big File

Consider the file **data/exercises/world_percentage_of_literacy.tsv**. Let's imagine this file is really huge, with a lot of rows!  Read the file using a random sample of 10 rows. Count the actual lines with `wc` 

<!-- grade_id: cell-fb39107093dd73fb, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/world_percentage_of_literacy.tsv with wc and save the number of lines
# in lines_in_file (as an integer), using 
lines_in_file = ! wc -l < data/exercises/world_percentage_of_literacy.tsv
# Notice that the header is not a line..
lines_in_file = int(lines_in_file[0])-1
lines_in_file
# YOUR CODE HERE
#raise NotImplementedError()

# make parameter rows_to_skip equal to the lines you want to skip loading 
# don't forget: 10 rows should be fecthed
sample_number = 10
n_rows_to_skip = lines_in_file - sample_number
# rows_to_skip = ...
random.seed(42) # this is to get always the same sample. can be removed if we want the sample to change
rows_to_skip = random.sample(
    range(1, lines_in_file-1), # this is a range from the first row after the header, to the last row on the file
    n_rows_to_skip # this is the number of rows we want to random sample here, and that we will be skipped on pd.read_csv with argument skiprows
)
# YOUR CODE HERE
#raise NotImplementedError()
# Create a df8 dataframe with the sampled values
# don't forget: you want to keep the header plus 10 rows in the new dataframe.
# YOUR CODE HERE
#raise NotImplementedError()


df8= pd.read_csv('data/exercises/world_percentage_of_literacy.tsv',sep='\t',
    skiprows=rows_to_skip
)
```

```python
df8['Gender difference'].str.replace('%', '').astype('float64').mean()
```

Output:
```
7.885714285714286
```

```python
df8.dropna()
```

Output:
```
Country Literacy rate (all) Male literacy Female literacy  \
0      Albania               97.6%         98.4%           96.8%   
2      Bahrain               95.7%         96.9%           93.5%   
3  Ivory Coast               43.1%         53.1%           32.5%   
4         Iran               86.8%         91.2%           82.5%   
5       Serbia               98.1%         99.1%           97.2%   
8       Zambia               63.4%         70.9%           56.0%   
9     Zimbabwe               86.5%         88.5%           84.6%   

  Gender difference  
0              1.6%  
2              3.5%  
3             20.6%  
4              8.7%  
5              1.9%  
8             14.9%  
9              4.0%
```

```python
df8['index']=range(0,len(df8.index))
```

```python
df8.set_index('index')
```

Output:
```
Country          Literacy rate (all) Male literacy Female literacy  \
index                                                                           
0          Albania                        97.6%         98.4%           96.8%   
1          Bahamas  not reported by UNESCO 2015           NaN             NaN   
2          Bahrain                        95.7%         96.9%           93.5%   
3      Ivory Coast                        43.1%         53.1%           32.5%   
4             Iran                        86.8%         91.2%           82.5%   
5           Serbia                        98.1%         99.1%           97.2%   
6           Sweden  not reported by UNESCO 2015           NaN             NaN   
7      Switzerland  not reported by UNESCO 2015           NaN             NaN   
8           Zambia                        63.4%         70.9%           56.0%   
9         Zimbabwe                        86.5%         88.5%           84.6%   

      Gender difference  
index       
```

```python
! head -11 data/exercises/world_percentage_of_literacy.tsv
```

Output:
```
Country	Literacy rate (all)	Male literacy	Female literacy	Gender difference
World	86.3%	90.0%	82.7%	7.3%
Afghanistan	38.2%	52.0%	24.2%	27.8%
Albania	97.6%	98.4%	96.8%	1.6%
Algeria	80.2%	87.2%	73.1%	14.0%
Andorra	not reported by UNESCO 2015			
Angola	71.1%	82.0%	60.7%	21.3%
Antigua and Barbuda	not reported by UNESCO 2015	99.0%(2013)[3][note 1]		
Argentina	98.1%	98.0%	98.1%	-0.1%
Armenia	99.8%	99.8%	99.7%	0.1%
Australia	not reported by UNESCO 2015
```

<!-- grade_id: cell-ec15a8be42137a01, points: 2, grade: true, solution: false -->
```python
# getting remaining list of countries, removing rows_to_skip from the file
df_helper = pd.read_csv( 
    'data/exercises/world_percentage_of_literacy.tsv', 
    sep='\t',
    header=0 
)

total_indexes = list(df_helper.index)
for s in rows_to_skip:
    total_indexes.remove(s-1)

expected_hash = '7559ca4a957c8c82ba04781cd66a68d6022229fca0e8e88d8e487c96ee4446d0'
assert hashlib.sha256(str(lines_in_file).encode()).hexdigest() == expected_hash, "lines_in_file are wrong"
assert isinstance(df8, pd.DataFrame)
assert list(df8['Country'])==list(df_helper.iloc[total_indexes]['Country'])
assert df8.shape[0]==10, "df8 size is wrong"

expected_hash = '4a44dc15364204a80fe80e9039455cc1608281820fe2b24f1e5233ade6af1dd5'
assert  hashlib.sha256(str(df8.shape[0]).encode()).hexdigest() == expected_hash

```

## Q9: Loading a Big File

Read file **data/exercises/world_percentage_of_literacy.tsv** using chunks keep only the columns `Country` and `Gender difference`.
Note that:
* file should be read by chunks of 20 countries
* the missing values should be removed (filtered in each chunk)
* the `Gender difference` should be converted to type float (in each chunk)
* the index should be incremental starting from 0 (i.e, you don't need to read any column as the index)

At the end, calculate the average `Gender difference`

Tip: Be sure you investigate the data you are about to load before doing any code.

```python
chunks_iter = pd.read_csv('data/exercises/world_percentage_of_literacy.tsv', sep ='\t',
    chunksize=20
)

for data_chunk in chunks_iter:
    data_chunk2=data_chunk.dropna().reset_index(drop=True)
    print(data_chunk2)
```

Output:
```
Country Literacy rate (all) Male literacy Female literacy  \
0         World               86.3%         90.0%           82.7%   
1   Afghanistan               38.2%         52.0%           24.2%   
2       Albania               97.6%         98.4%           96.8%   
3       Algeria               80.2%         87.2%           73.1%   
4        Angola               71.1%         82.0%           60.7%   
5     Argentina               98.1%         98.0%           98.1%   
6       Armenia               99.8%         99.8%           99.7%   
7    Azerbaijan               99.8%         99.9%           99.7%   
8       Bahrain               95.7%         96.9%           93.5%   
9    Bangladesh       72.8%[note 2]         75.6%           69.9%   
10      Belarus               99.7%         99.8%           99.7%   
11       Belize               82.7%         82.3%           83.0%   
12        Benin               38.4%         49.9%           27.3%   

   Gender difference  
0               7.
```

<!-- grade_id: cell-5106c25a705296e8, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/world_percentage_of_literacy.tsv
# the chunks should be appended in a list called chunk_arr
# YOUR CODE HERE
#raise NotImplementedError()
chunksize=20
chunks_iter = pd.read_csv('data/exercises/world_percentage_of_literacy.tsv', sep ='\t', header=0,
    chunksize=20
)

#indice=0
chunk_arr = []
for data_chunk in chunks_iter:
    data_chunk_filtered = data_chunk.dropna()
    data_chunk_filtered=data_chunk_filtered.reset_index(drop=True)
    data_chunk_filtered['Gender difference']=data_chunk_filtered['Gender difference'].str.replace('%', '').astype('float64')
    #data_chunk_filtered=data_chunk_filtered.reindex(index_)
    #indice += len(data_chunk_filtered.index) 
    chunk_arr.append(data_chunk_filtered)
    
df9 = pd.concat(chunk_arr, axis=0)
#index_=list(range(len(df9.index)))
df9=df9.reset_index(drop=True)
#df9.head()

# df9 should be the final dataframe with concatenated chunks
# Resulting average should go on lit_avg variable
lit_avg=df9['Gender difference'].mean()
# YOUR CODE HERE
#raise NotImplementedError()
```

```python
! head -10 data/exercises/world_percentage_of_literacy.tsv
```

Output:
```
Country	Literacy rate (all)	Male literacy	Female literacy	Gender difference
World	86.3%	90.0%	82.7%	7.3%
Afghanistan	38.2%	52.0%	24.2%	27.8%
Albania	97.6%	98.4%	96.8%	1.6%
Algeria	80.2%	87.2%	73.1%	14.0%
Andorra	not reported by UNESCO 2015			
Angola	71.1%	82.0%	60.7%	21.3%
Antigua and Barbuda	not reported by UNESCO 2015	99.0%(2013)[3][note 1]		
Argentina	98.1%	98.0%	98.1%	-0.1%
Armenia	99.8%	99.8%	99.7%	0.1%
```

```python
df9['Literacy rate (all)'].apply(lambda x: x.split('%')[0]).astype('float64')
```

Output:
```
0      86.3
1      38.2
2      97.6
3      80.2
4      71.1
       ... 
146    95.4
147    94.5
148    70.1
149    63.4
150    86.5
Name: Literacy rate (all), Length: 151, dtype: float64
```

<!-- grade_id: cell-f60a22c3465d5b83, points: 2, grade: true, solution: false -->
```python
expected_hash = 'cc63c8d6f0f9a8260d40b5c9a62d4fce5b693bab5750c90fd1f038ff64e37d6d'
assert hashlib.sha256(str([len(c) for c in chunk_arr]).encode()).hexdigest() == expected_hash, "error on chunk_arr sizes"

assert df9.loc[df9['Country']=='World', 'Gender difference'].values[0] == 7.3, "df9 data is wrong"
assert df9.dtypes['Country'] == object, "df9 structure is wrong"
assert df9.dtypes['Gender difference'] == float, "df9 structure is wrong" 

expected_hash = 'c4d7d45402099baaea8a079580a4ef950016ef7403178f72e64ceb3738e4471d'
assert hashlib.sha256(str(lit_avg).encode()).hexdigest() == expected_hash, "lit_avg is wrong"
```

## Q10: Calculate average values of `Literacy All` using chunks and avoiding a complete data frame in memory.

Using chunks, read file **data/exercises/world_percentage_of_literacy.tsv**, avoid incomplete rows and calculate the average of ***Literacy rate (all)*** without loading all data simultaneously. 

Use a similar approach of the previous question but don't create any dataframe neither any list with chunks; ***Hint: Use the average definition***



<!-- grade_id: cell-57fd7700885246eb, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/world_percentage_of_literacy.tsv
# the final average should be in the variable final_avg
# You should increment 2 variables in each chunk and use them at the end to calculate final_avg. call them lit_a, lit_b
# YOUR CODE HERE
#raise NotImplementedError()
#chunksize=20
chunks_iter = pd.read_csv('data/exercises/world_percentage_of_literacy.tsv', sep ='\t', header=0,
    chunksize=2
)

lit_a=0
lit_b=0

chunk_arr = []
for data_chunk in chunks_iter:
    data_chunk_filtered = data_chunk.dropna()
    data_chunk_filtered=data_chunk_filtered.reset_index(drop=True)
    data_chunk_filtered['Literacy rate (all)']=data_chunk_filtered['Literacy rate (all)'].apply(lambda x: x.split('%')[0]).astype('float64')

    lit_a += data_chunk_filtered['Literacy rate (all)'].sum()

    lit_b += len(data_chunk_filtered['Literacy rate (all)'])

    

print(lit_a)
print(lit_b)


final_avg=(lit_a/lit_b)
# YOUR CODE HERE
#raise NotImplementedError()
print(final_avg)

```

Output:
```
12569.8
151
83.24370860927152
```

<!-- grade_id: cell-0edcf1260d7886c9, points: 2, grade: true, solution: false -->
```python
assert math.isclose(83.24370860927154, final_avg, rel_tol=1e-1), "final_avg is wrong"
assert lit_b==151 or lit_a==151, "lit_a or lit_b is wrong"
assert int(lit_b) == 12569 or int(lit_a)==12569, "lit_a or lit_b is wrong"
```

```python

```